# Problem Set 4

**Q1.** Answer the following questions:

- How to create a tuple? Say with elements 3, 4, '3', and '4'.
- How to create a tuple from a list? Say from `[3, 4, '3', '4']`.
- How to create a tuple from a string? Say from `"abac"`.
- How to get a length of a tuple?
- How to get the n-th element of a tuple? 
- How to combine or concatenate two tuples?
- How to duplicate a tuple?
- How to create a tuple from another tuple by slicing?
- How to check if an object is an element of a tuple?
- How to traverse a tuple?
- What's the difference between tuple and list?

```python
# Try it here

###
### YOUR CODE HERE
###

```

**Q2.** Answer the following questions:

- How do you create a set?
- How to create a set from a list or a tuple?
- How to create a set from a string?
- How to add item into a set?
- How to find the number of item in the set?
- How to remove an item from the set?
- How to check if a set a subset of another set?
- How to check if a set is a superset of another set?
- How to check if a set is a proper subset or a proper superset of another set?
- How to create a union from two sets?
- How to create an intersection from two sets?
- How to find the difference between set 1 and set 2? 
- How to find a symmetric difference of two sets?

```python
# Try it here

###
### YOUR CODE HERE
###

```

**Q3.** Answer the following Questions
- How to create a dictionary?
- How to get a value given a key in a dictionary?
- How to add a key-value pair into a dictionary?
- How to modify a key-value pair in a dictionary?
- How to delete a key-value pair from a dictionary?
- How to find the number of key-value pairs in a dictionary?
- How to test if a key is in a dictionary?
- How to get all the keys in a dictionary?
- How to get all the values in a dictionary?
- How to traverse a dictionary?
- What's the difference between a dictionary and a set?


```python
# try it here

###
### YOUR CODE HERE
###

```

**Q4.** *Counting Keywords:* Write a function that takes in two arguments. The first argument is a *set* of keywords, and the second arugment is a string of Python codes. The function should return how many keywords are there in the code.

```python
def count_keywords(keyword, code):
    pass
    ###
    ### YOUR CODE HERE
    ###
    

```

<!-- grade_id: Q4, points: 1, grade: true, solution: false -->
```python
keywords = {"and", "as", "assert", "break", "class",
            "continue", "def", "del", "elif", "else",
            "except", "False", "finally", "for", "from",
            "global", "if", "import", "in", "is", "lambda",
            "None", "nonlocal", "not", "or", "pass", "raise",
            "return", "True", "try", "while", "with", "yield"}

inp1 = """
def accumulate(m):
    total = 0
    for row in m:
        total += sum(row)
    return total
"""

inp2 = """
import random
def get_matrix(rows, cols):
    matrix = []
    for row in range(rows):
        matrix.append([])
        for col in range(cols):
            value = random.randint(0,100)
            matrix[row].append(value)
    return matrix
"""

assert count_keywords(keywords, inp1) == 4
assert count_keywords(keywords, inp2) == 7
```

**Q5.** Modify Q4 so that it returns not only the number of keywords but a tuple of all the keywords found. If there are more than one keywords, those duplicates should be in the tuple in the sequence that you found them from the first to the last.

```python
def count_find_keywords(keywords, code):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q5, points: 1, grade: true, solution: false -->
```python
keywords = {"and", "as", "assert", "break", "class",
            "continue", "def", "del", "elif", "else",
            "except", "False", "finally", "for", "from",
            "global", "if", "import", "in", "is", "lambda",
            "None", "nonlocal", "not", "or", "pass", "raise",
            "return", "True", "try", "while", "with", "yield"}

inp1 = """
def accumulate(m):
    total = 0
    for row in m:
        total += sum(row)
    return total
"""

inp2 = """
import random
def get_matrix(rows, cols):
    matrix = []
    for row in range(rows):
        matrix.append([])
        for col in range(cols):
            value = random.randint(0,100)
            matrix[row].append(value)
    return matrix
"""

assert count_find_keywords(keywords, inp1) == (4, ("def", "for", "in", "return"))
assert count_find_keywords(keywords, inp2) == (7, ("import", "def", "for", "in", "for", "in", "return"))
```

**Q6.** *counting occurence of keywords:* Modify Q5 so that it return values is a dictionary where the keys are the list of unique keywords found and its values is the number of occurences of that keywrod.

```python
def extract_keywords(keywords, code):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q6, points: 1, grade: true, solution: false -->
```python
keywords = {"and", "as", "assert", "break", "class",
            "continue", "def", "del", "elif", "else",
            "except", "False", "finally", "for", "from",
            "global", "if", "import", "in", "is", "lambda",
            "None", "nonlocal", "not", "or", "pass", "raise",
            "return", "True", "try", "while", "with", "yield"}

inp1 = """
def accumulate(m):
    total = 0
    for row in m:
        total += sum(row)
    return total
"""

inp2 = """
import random
def get_matrix(rows, cols):
    matrix = []
    for row in range(rows):
        matrix.append([])
        for col in range(cols):
            value = random.randint(0,100)
            matrix[row].append(value)
    return matrix
"""

assert extract_keywords(keywords, inp1) == {"def": 1, "for": 1, "in": 1, "return":1}
assert extract_keywords(keywords, inp2) == {"import": 1, "def": 1, "for": 2, "in": 2, "return": 1}
```

**Q7.** *Hangman:* This question and the followings are part of a problem to create a game called Hangman. Write a function that takes in a list of tuples and returns a dictionary. The input is a list of tuples where the element of the tuple is in the form `(category, word)`. For example, we have the following list:
```
[('animal', 'cat'), ('animal', 'dog'), ('food', 'rice'), ...]
```

The function should return a dictinary where the keys are the categories and the values are the set of words in that category. For example the above input list will result in:
```
{'animal': {'cat', 'dog', ...}, 'food': {'rice', ...}}
```

```python
def create_dict(list_words):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q7, points: 1, grade: true, solution: false -->
```python
words = [('food', 'apple'), ('food', 'almond'), ('food', 'apricot'), ('food', 'beancurd'), ('food', 'bacon'),
         ('food', 'bread'), ('food', 'beans'), ('food', 'broccoli'), ('food', 'cake'), ('food', 'chili'),
         ('animal', 'ape'), ('animal', 'ant'), ('animal', 'bat'), ('animal', 'bear'), ('animal', 'cat'),
         ('animal', 'dog'), ('animal', 'flea'), ('animal', 'horse'), ('animal', 'lamb'), ('animal', 'leopard'),
         ('job', 'baker'), ('job', 'farmer'), ('job', 'builder'), ('job', 'teacher'), ('job', 'pilot'), ('job', 'gardener'),
         ('job', 'gardener'), ('job', 'secretary'), ('job', 'salesman'), ('job', 'programmer'), ('job', 'manager')]

assert create_dict(words) == {'food': {'apricot', 'cake', 'bacon', 'chili', 'almond', 'beancurd', 'apple', 'bread', 'broccoli', 'beans'}, 
                              'animal': {'dog', 'bat', 'horse', 'leopard', 'bear', 'ape', 'cat', 'flea', 'lamb', 'ant'}, 
                              'job': {'pilot', 'programmer', 'manager', 'gardener', 'farmer', 'teacher', 'builder', 'baker', 'secretary', 'salesman'}}
```

**Q8.** *Hangman:* Write a function to get a random word from a dictionary of words given a category. The function should take in three arguments. The first argument is the category to get the random word from and the second argument is the dictionary of the words. The third argument is optional, which is the seed number for the random module. By default the seed argument is set to 100. *Hint: use the random module.*

```python
import random

def get_random_word(category, word_dict, seed=100):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

```python
words = {'food': {'apricot', 'cake', 'bacon', 'chili', 'almond', 'beancurd', 'apple', 'bread', 'broccoli', 'beans'}, 
                              'animal': {'dog', 'bat', 'horse', 'leopard', 'bear', 'ape', 'cat', 'flea', 'lamb', 'ant'}, 
                              'job': {'pilot', 'programmer', 'manager', 'gardener', 'farmer', 'teacher', 'builder', 'baker', 'secretary', 'salesman'}}


print(get_random_word('food', words))
print(get_random_word('animal', words))
print(get_random_word('animal', words, 200))
print(get_random_word('job', words))
```

**Q9.** *Hangman:* Write a function `hidden_word(word)` that takes in a string argument which is the chosen word and returns a tuple. The first element of the tuple is a string of asterix with the same length of the word and the second element is the word itself. For example:
```
hidden_word('bacon')
```
should return a tuple:
```
('*****', 'bacon')
```


```python
def hidden_word(word):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q9, points: 1, grade: true, solution: false -->
```python
assert hidden_word('bacon') == ('*****', 'bacon')
assert hidden_word('programmer') == ('**********', 'programmer')
assert hidden_word('dog') == ('***', 'dog')
```

**Q10.** *Hangman:* Write a function `check_letter(hidden_word, guess)` which takes in two arguments where the first argument is a tuple in the form of the output of the previous question and the second argument is a single character string. The function should check if the guess character is in the word. If it is, it should return a tuple where the first argument is `True` and the second argument is a tuple of the `hidden_word` with the asterix replaced with the guess words. For example,

```
check_letter(('**********', 'programmer'), 'r')
```
should return:
```
(True, ('*r**r****r', 'programmer')))
```
If the guess letter is not in the word, it should return `False` and the tuple unchanged. For example,
```
check_letter(('**********', 'programmer'), 'q')
```
should return:
```
(False, ('**********', 'programmer')))
```


```python
def check_letter(hidden_word, guess):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q10, points: 1, grade: true, solution: false -->
```python
assert check_letter(('**********', 'programmer'), 'r') == (True, ('*r**r****r', 'programmer'))
assert check_letter(('*r**r****r', 'programmer'), 'm') == (True, ('*r**r*mm*r', 'programmer'))
assert check_letter(('*r**r*mm*r', 'programmer'), 'p') == (True, ('pr**r*mm*r', 'programmer'))

```