# Problem Set 10

**Q1.** *Inheritance:* Create a class called `MixedFraction` as a subclass of `Fraction`. A mixed fraction is a fraction that comprises of a whole number, a numerator and a denominator, e.g. `1 2/3` which is the same as `5/3`. The class has the following way of initializing its properties:
- `__init__(top, bot, whole)`: which takes in three Integers, the whole number, the numerator, and the denominator, e.g. `whole=1`, `top=2`, `bot=3`. The argument `whole` by default is `0`.  You can also specify `top` to be greater than `bot`. 

The class only has two properties:
- `num`: which is the numerator and can be greater than denominator.
- `den`: which is the denominator and must be a non-zero number.

The class should also have the following methods:
- `get_three_numbers()`: which is used to calculate the whole number, numerator and the denominator from a given numerator and denominator. The stored properties are `num` and `den` as in `Fraction` class. This function returns three Integers as a tuple, i.e. `(whole, top, bot)`.

The class should also override the `__str__()` method in this manner:
- `num/dem` if the numerator is smaller than the denominator. For example, `2/3`. 
- `whole top/bot` if the numerator is greater than the denominator. For example, `1 2/3`.

```python
def gcd(a, b):
    ###
    ### YOUR CODE HERE
    ###
    pass


class Fraction:
    ###
    ### YOUR CODE HERE
    ###
    pass
```

```python
class MixedFraction(Fraction):
  def __init__(self, top, bot, whole=0):
    ###
    ### YOUR CODE HERE
    ###
    pass

  def get_three_numbers(self):
    ###
    ### YOUR CODE HERE
    ###
    return (whole, top, bot)

  def __add__(self, other):
    ###
    ### YOUR CODE HERE
    ###

  def __str__(self):
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q1a, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (1, 2, 3)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
print(type(result))
assert result.num == 10 and result.den == 3

assert mf1 == mf2
```

Output:
```
<class '__main__.MixedFraction'>
```

<!-- grade_id: Q1b, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (1, 2, 3)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
assert result.num == 10 and result.den == 3

result = mf1 * mf2
assert result.num == 25 and result.den == 9

mf3 = MixedFraction(1, 2, 1)
result = mf1 - mf3
assert result.num == 1 and result.den == 6

assert str(mf1) == "1 2/3"
```

**Q2.** Write a class called `EvaluateFraction` that evaluates postfix notation implemented Stack and Queue data structures. Postfix notation is a way of writing expressions without using parenthesis. For example, the expression `(1+2)*3` would be written as `1 2 + 3 *`. The class `EvaluateFraction` has the following method:
- `input(inp)`: which pushes the input input one at a time. For example, to create a postfix notation `1 2 + 3 *`, we can call this method repetitively, e.g. `e.input('1'); e.input('2'); e.input('+'); e.input('3'); e.input('*')`. Notice that the input is of String data type. 
- `evaluate()`: which returns the output of the expression.
- `get_fraction(inp)`: which takes in an input string and returns a `Fraction` object. 

Postfix notation is evaluated using a Stack. You can implement Stack using your previous Stack class. The input streams from `input()` are stored in a Queue, which we will again implement using your previous Queue class. If the output of the Queue is a number, the item is pushed into the stack. If it is an operator, we will apply the operator to the two top most item n the stacks and push the result back into the stack. 

```python
class Stack:
    ###
    ### YOUR CODE HERE
    ###
    pass
```

```python
#class Queue:
#    ### BEGIN SOLUTION
#     def __init__(self):
#         self._data = []

#     def put(self, item):
#         self._data.append(item)

#     def get(self):
#         return self._data.pop(0)

#     def empty(self):
#         return len(self._data) == 0

#     def qsize(self):
#         return len(self._data)

class Queue:
    def __init__(self):
        self._data = []

    def put(self, item):
        self._data.append(item)

    def get(self):
        return self._data.pop(0)

    @property
    def empty(self):
        return len(self._data) == 0
    @property
    def qsize(self):
        return len(self._data)
    ### END SOLUTION
    pass
```

```python
class EvaluateFraction:

  operands = "0123456789"
  operators = "+-*/"

  def __init__(self):
    self.expression = Queue()
    self.stack = Stack()

  def input(self, item):
    ###
    ### YOUR CODE HERE
    ###
    pass

  def evaluate(self):
    ###
    ### YOUR CODE HERE
    ###
    pass

  def get_fraction(self, inp):
    pass
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q2, points: 1, grade: true, solution: false -->
```python
pe = EvaluateFraction()
pe.input("1/2")
pe.input("2/3")
pe.input("+")
assert pe.evaluate()==Fraction(7, 6)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
assert pe.evaluate()==Fraction(1, 1)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1/6")
pe.input("-")
pe.input("3/4")
pe.input("*")
assert pe.evaluate()==Fraction(3, 4)
```

**Q3.** Modify Q2 so that it can work with MixedFraction. Write a class called `EvaluateMixedFraction` as a subclass of `EvaluateFraction`. You need to override the following methods:
- `get_fraction(inp)`: This function should be able to handle string input for MixedFraction such as `1 1/2` or `3/2`. It should return a `MixedFraction` object.
- `evaluate()`: This function should return `MixedFraction` object rather than `Fraction` object. 

```python
class EvaluateMixedFraction(EvaluateFraction):
  ###
  ### YOUR CODE HERE
  ###
  pass
    
```

<!-- grade_id: Q3, points: 1, grade: true, solution: false -->
```python
pe = EvaluateMixedFraction()
pe.input("3/2")
pe.input("1 2/3")
pe.input("+")
assert pe.evaluate() == MixedFraction(1, 6, 3)

pe.input("1/2")
pe.input("2/3")
pe.input("+")
pe.input("1 1/8")
pe.input("-")
assert pe.evaluate() == MixedFraction(1, 24)

pe.input("1 1/2")
pe.input("2 2/3")
pe.input("+")
pe.input("1 1/6")
pe.input("-")
pe.input("5/4")
pe.input("*")
assert pe.evaluate() == MixedFraction( 3, 4, 3)
```

## Vehicle Rental Problem

### The Problem

The problem is to develop an object-oriented design and implementation of a program capable of maintaining reservations for a vehicle rental agency. The agency rents out three types of vehicles: 
- cars, 
- vans, and 
- moving trucks. 

The program should allow users to:
- check for available vehicles
- request rental charges by vehicle type
- get the cost of renting a particular type vehicle for a specified period of time, and
- make/cancel reservations

See the full problem in: [Vehicle Rental Problem](https://sutdapac-my.sharepoint.com/:b:/g/personal/oka_kurniawan_sutd_edu_sg/EQeebPYHPUBOo7PIeOJDKrsBURo53xBvwBIvi-lq2ToLCA?e=UXgblf)

**Q4.** *Classes: Vehicle Polymorphism.* Create a super class Vehicle that takes in miles per gallon (mpg) and vehicle identification number (vin). The class should support a method to return the type of the Vehicle. There are three types: Car, Van, and Truck. It should also supports a method to get VIN, and description. The description method should return a string that states the attributes that are common to all vehicles, i.e. mpg and vin. It should also support a method to check if the vehicle is reserved or not, and a method to set the reservation status. The class Vehicle has three sub classes, i.e. Car, Van and Truck. Each of these classes have different arguments for its initialization. The method that gives the description is overridden in these child classes to give a more detailed description of the class. 

```python
class Vehicle:
    def __init__(self, mpg, vin):
        ###
        ### YOUR CODE HERE
        ###
        pass

    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q4, points: 1, grade: true, solution: false -->
```python

v = Vehicle(12, 'EJ5KU2435BC')
assert v._mpg == 12 and v._vin == 'EJ5KU2435BC'

c = Car('Chevrolet Camaro', 30, 4, 2, 'WG8JM5492DY')
vn = Van('Chrysler Town&Country', 25, 7, 'DK3KG8312UE')
t = Truck(12, 'Ten-Foot', 1, 'EJ5KU2435BC')

assert c.get_type() == 'Car' and vn.get_type() == 'Van' and t.get_type() == 'Truck'

assert c.vin == 'WG8JM5492DY'

assert not vn.reserved 

t.reserved = True
assert t.reserved 

t.reserved = 'False'
assert t.reserved

t.reserved = False

assert not t.reserved

print(c.get_description())
print(vn.get_description())
print(t.get_description())
```

Output:
```
Chevrolet Camaro     passengers:  4   doors:  2   mpg:  30   vin:  WG8JM5492DY
Chrysler Town&Country    passengers:  7   mpg:  25   vin:  DK3KG8312UE
length(feet): Ten-Foot   rooms:  1   mpg:  12   vin:  EJ5KU2435BC
```

**Q5.** *Classes: Vehicle Rental Program. (optional)* Refer to the other PDF handout and complete the other parts of Vehicle Rental Program.

Hints:
- Refer to the UML diagram below.
- Use [PEP8](https://www.python.org/dev/peps/pep-0008/) convention in the namings.

![UML Diagram](https://www.dropbox.com/s/88tpf6uds03n3li/vehicle_uml.png?raw=1)

```python

```