# Problem Set 5

**Q1.** *Sorting: Insertion Sort.* Write a function to do insertion sort for a list of numbers. Check: [Wikipedia's article on Insertion Sort](https://en.wikipedia.org/wiki/Insertion_sort). Measure the computation for a random list with size: 10, 100, 1000, 10 000, and 100 000. Store the time taken in a list and plot a graph of time taken versus the number of items in the list. 

```python
import random
import timeit
import copy
import math

def insertion_sort(list_items):
    ###
    ### YOUR CODE HERE
    ###
    pass


```

<!-- grade_id: Q1, points: 2, grade: true, solution: false -->
```python

assert insertion_sort([54, 26, 93, 17, 77, 31, 44, 55, 20]) == [17, 20, 26, 31, 44, 54, 55, 77, 93]

```

```python

preparation_insertion = """
import random
from __main__ import insertion_sort
"""


data ="""
items = list(range({}))
random.shuffle(items)
"""

preparation_sort = """
import random
"""
x_axis = [10**n for n in range(1,5)]
data_list = [data.format(x) for x in x_axis]

insertion_time = []

for x in range(len(x_axis)):
	t = timeit.Timer("insertion_sort(items)",preparation_insertion + data_list[x])
	insertion_time.append(t.timeit(number=1))

print(insertion_time)


```

Output:
```
[8.499999999855845e-06, 0.00033839999999685233, 0.048735700000001714, 3.8915729999999975]
```

```python

import matplotlib.pyplot as plt
plt.subplot(2,1,1)
plt.plot(x_axis,insertion_time,'o-')

plt.subplot(2,1,2)
x2_axis = [x*x for x in x_axis]
plt.plot(x2_axis,insertion_time,'o-')

plt.show()
```

Output:
```
<Figure size 640x480 with 2 Axes>
```

**Q2.** *Computation Time* Do the same time measurement using Python's built-in function, i.e. `sorted()`. Plot a graph of time taken versus the number of items in the list.

```python
import random
import timeit
import copy
import math

preparation_insertion = """
import random
"""


data ="""
items = list(range({}))
random.shuffle(items)
"""

preparation_sort = """
import random
"""
x_axis = [10**n for n in range(1,5)]
data_list = [data.format(x) for x in x_axis]

sort_time = []

for x in range(len(x_axis)):
	t = timeit.Timer("sorted(items)",preparation_sort + data_list[x])
	sort_time.append(t.timeit(number=1))

print(sort_time)


```

Output:
```
[1.900000000887303e-06, 6.8000000013057615e-06, 8.740000000173609e-05, 0.0018689999999992324]
```

```python

plt.plot(x_axis,sort_time,'o-')
plt.show()
```

Output:
```
<Figure size 432x288 with 1 Axes>
```

**Q3.** *Asymptotic Notation.* Modify the x-axis so that you can get a linear graph for both insertion sort and Python's sort computation time. 


```python

xlogx_axis =[x*math.log(x) for x in x_axis]
plt.plot(xlogx_axis,sort_time,'o-')
plt.show()
```

Output:
```
<Figure size 432x288 with 1 Axes>
```