# Problem Set 7

**Q1.** *Recursion: Palindrome.* Redo palindrome problem with Recursion. Derive the recurrence relations and analyze its complexity. 

```python
def palindrome(string):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: Q1, points: 1, grade: true, solution: false -->
```python
assert palindrome('a')
assert palindrome('bb')
assert not palindrome('ab')
assert not palindrome('abc')
assert palindrome('aba')
assert palindrome('123454321')
```

**Q2.** *Recursion: Coin Change.* Write a function that find the smallest number of coins that represent the given value. There are only four possible coins: 10 cents, 20 cents, 50 cents, and 100 cents (1 dollar). The function should return a dictionary with the coin value as its key and the number of coins as its value. For example:

```python
>>> assert coin_change(50) == {50:1}

>>> assert coin_change(80) == {50:1, 20:1, 10:1}

>>> assert coin_change(120) == {100:1, 20:1}

>>> assert coin_change(90) == {50:1, 20:2}
```

```python
coin_list = (100, 50, 20, 10)

def coin_change(total_sum):
    ###
    ### YOUR CODE HERE
    ###
    pass

```

<!-- grade_id: Q2, points: 1, grade: true, solution: false -->
```python
assert coin_change(50) == {50:1}
assert coin_change(80) == {50:1, 20:1, 10:1}
assert coin_change(120) == {100:1, 20:1}
assert coin_change(90) == {50:1, 20:2}
assert coin_change(40) == {20:2}

try:
    coin_change(31) 
except ValueError as e:
    print('Error:',e)

```

Output:
```
Error: No possible coin can be found for 1
```

**Q3.** *Towers of Hanoi:* Write a function `move_disks(n, from_tower, to_tower, aux_tower)` which returns an array of String for the movement of disks that solves the [Towers of Hanoi problem](https://en.wikipedia.org/wiki/Tower_of_Hanoi). 
* The first argument `n` is an Integer input that gives information on the number of disk. 
* The second argument `from_tower` is a String which is the label of the origin tower.
* The third argument `to_tower` is a String which is the label of the destination tower.
* The last argument `aux_tower` is a String which is the label of the auxilary tower.


```python
def move_disks(n, from_tower, to_tower, aux_tower):
  result = []
  ###
  ### YOUR CODE HERE
  ###
  pass

  return result


```

<!-- grade_id: Q3, points: 1, grade: true, solution: false -->
```python
result = move_disks(3, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to B.", "Move disk 2 from A to C.", "Move disk 1 from B to C.", "Move disk 3 from A to B.", "Move disk 1 from C to A.", "Move disk 2 from C to B.", "Move disk 1 from A to B."]

result = move_disks(4, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B.", "Move disk 3 from A to C.", "Move disk 1 from B to A.", "Move disk 2 from B to C.", "Move disk 1 from A to C.", "Move disk 4 from A to B.", "Move disk 1 from C to B.", "Move disk 2 from C to A.", "Move disk 1 from B to A.", "Move disk 3 from C to B.", "Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B."]
result = move_disks(5, "A", "B", "C")
assert result == ['Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.', 'Move disk 4 from A to C.', 'Move disk 1 from B to C.', 'Move disk 2 from B to A.', 'Move disk 1 from C to A.', 'Move disk 3 from B to C.', 'Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 5 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.', 'Move disk 3 from C to A.', 'Move disk 1 from B to C.', 'Move disk 2 from B to A.', 'Move disk 1 from C to A.', 'Move disk 4 from C to B.', 'Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.']
```

Output:
```
['Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.']
['Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.', 'Move disk 3 from A to C.', 'Move disk 1 from B to A.', 'Move disk 2 from B to C.', 'Move disk 1 from A to C.', 'Move disk 4 from A to B.', 'Move disk 1 from C to B.', 'Move disk 2 from C to A.', 'Move disk 1 from B to A.', 'Move disk 3 from C to B.', 'Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.']
```

**Q4.** *Recursion: Merge Sort.* Write a function that does Merge sort on a list of numbers. Measure the time and compare its computation time with that of insertion sort and Python's built-in function `sorted()`.


```python
import random
import timeit
import copy
import math

def merge_sort(items):
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: Q4, points: 1, grade: true, solution: false -->
```python

list_a = [54, 26, 93, 17, 77, 31, 44, 55, 20]
merge_sort(list_a)
assert  list_a == [17, 20, 26, 31, 44, 54, 55, 77, 93]


```

```python
preparation_insertion = """
import random
from __main__ import merge_sort
"""


data ="""
items = list(range({}))
random.shuffle(items)
"""

preparation_sort = """
import random
"""
x_axis = [10**n for n in range(1,5)]
data_list = [data.format(x) for x in x_axis]

merge_time = []
sort_time = []

for x in range(len(x_axis)):
	t = timeit.Timer("merge_sort(items)",preparation_insertion + data_list[x])
	merge_time.append(t.timeit(number=1))
	t = timeit.Timer("sorted(items)",preparation_sort + data_list[x])
	sort_time.append(t.timeit(number=1))

print(merge_time)
print(sort_time)

```

Output:
```
[1.8159999999989296e-05, 0.00019857099999998074, 0.002542218000000096, 0.03205288500000003]
[1.847000000054777e-06, 7.80000000000225e-06, 8.872399999992897e-05, 0.0011361280000000695]
```

```python
import matplotlib.pyplot as plt

plt.subplot(2,1,1)
xlogx_axis =[x*math.log(x) for x in x_axis]
plt.plot(xlogx_axis,merge_time,'o-')

plt.subplot(2,1,2)
plt.plot(xlogx_axis,sort_time,'o-')

plt.show()

```

Output:
```
<Figure size 432x288 with 2 Axes>
```

```python

```